---
title: TEST
description: 'Testing the standardization'
---

## 4.4 MB datasets

```yaml
type: NormalExercise
key: f4e4dacd43
lang: r
xp: 100
skills: 1
```

This exercise uses a 4.4 MB dataset.

# This is an h1 header

**this is bold**

here's something newwwwwww

`@instructions`
Preview the dataset.

`@hint`
Some hint.

`@pre_exercise_code`
```{r}
library(tidyverse)
billboard_eighth <- read_csv("https://s3.amazonaws.com/assets.datacamp.com/production/repositories/3140/datasets/billboard_eighth.csv")
```

`@sample_code`
```{r}
billboard_eighth
```

`@solution`
```{r}
billboard_eighth
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Nice!")
```

`@possible_answers`


`@feedback`


---

## 10.1 MB dataset

```yaml
type: NormalExercise
key: 48f0714614
lang: r
xp: 100
skills: 1
```

This exercise uses a 10.1 MB dataset.

`@instructions`
Preview the dataset.

`@hint`
Some hint.

`@pre_exercise_code`
```{r}
library(tidyverse)
billboard_quarter <- read_csv("https://s3.amazonaws.com/assets.datacamp.com/production/repositories/3140/datasets/billboard_quarter.csv")
```

`@sample_code`
```{r}
billboard_quarter
```

`@solution`
```{r}
billboard_quarter
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Nice!")
```

`@possible_answers`


`@feedback`


---

## 20.2 MB dataset

```yaml
type: NormalExercise
key: 90709f1dc0
lang: r
xp: 100
skills: 1
```

This exercise uses a 20.2 MB dataset.

`@instructions`
Preview the dataset.

`@hint`
Some hint.

`@pre_exercise_code`
```{r}
library(tidyverse)
billboard_half <- read_csv("https://s3.amazonaws.com/assets.datacamp.com/production/repositories/3140/datasets/billboard_half.csv")
```

`@sample_code`
```{r}
billboard_half
```

`@solution`
```{r}
billboard_half
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Nice!")
```

`@possible_answers`


`@feedback`


---

## 41.6 MB dataset

```yaml
type: NormalExercise
key: b0451dd812
lang: r
xp: 100
skills: 1
```



`@instructions`
Adding them back.

`@hint`
Some hint.

`@pre_exercise_code`
```{r}
library(tidyverse)
billboard_full <- read_csv("https://s3.amazonaws.com/assets.datacamp.com/production/repositories/3140/datasets/billboard_full.csv")
```

`@sample_code`
```{r}
billboard_full
```

`@solution`
```{r}
billboard_full
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Nice!")
```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: f627c8eb1a
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: ea0e7cd088
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 9e75c2ca75
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 0f287d6950
xp: 100
```



`@instructions`
test

`@hint`
test

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
